---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Exercise

```yaml
type: PureMultipleChoiceExercise
key: 86ec2c4253
xp: 50
```

What is my name?

`@hint`
Starts with L

`@possible_answers`
Lucy, [Lauren], Lisa, Laverne

`@feedback`
hi
